```python
# Initialize Otter
import otter
grader = otter.Notebook("Econ126_Homework_01.ipynb")
```

If the previous cell returns an error, then you need to install Otter-Grader by running the following command in a command prompt or in a Notebook cell by itself:

    pip install otter-grader

After installing, delete the cell where you ran the code to install Otter-Grader (click just to the right of the cell and type "x" or use the scissor button in the menu). Restart the Notebook using the refresh button in the menu.

# Homework 1 

**Instructions:** Complete the notebook below. Upload to completed file to the corresponding assignment on Gradescope.

## Exercise: Strings

Follow the instructions in the following cell. *Be sure to spell all variable names correctly.*

```python
# Create a variable called `city` that equals "irVinE". Print the value of the variable `city`.
# BEGIN SOLUTION
city = 'irVinE'
print(city)
# END SOLUTION

# Use the string method `title()` on the variable `city` to create a variable called `city_title`
#    that has its first character capitalized. Print the value of the variable `city_title`.
# BEGIN SOLUTION
city_title = city.title()
print(city_title)
# END SOLUTION

# Use the string method `upper()` on the variable `city` to create a variable called `city_upper`
#    that has all characters capitalized. Print the value of the variable `city_upper`.
# BEGIN SOLUTION
city_upper = city.upper()
print(city_upper)
# END SOLUTION

# Use the string method `swapcase()` on the variable `city` to create a variable called `city_swapped`
#    that reverses the case of each character in `the variable city`. Print the value of the variable 
#    `city_swapped`.
# BEGIN SOLUTION
city_swapped = city.swapcase()
print(city_swapped)
# END SOLUTION

# Use the string method `replace()` on the variable `city` to create a variable called `city_replaced`
#    that replaces all `i`s in the variable `city` by `a`s. Print the value of the variable
#    `city_replaced`.
# BEGIN SOLUTION
city_replaced = city.replace('i','a')
print(city_replaced)
# END SOLUTION

# Create a new variable called `state` equal to `CA`. Use `city` and `state` (and the `title()` method) 
#     to print `Irvine, CA`.
# BEGIN SOLUTION
state = 'CA'
print(city.title()+', '+state)
# END SOLUTION

# Use the string methods `join()` and `title()` to combine the variables `city` and `state` into 
#    `Irvine, CA` and print `Irvine, CA`.
# BEGIN SOLUTION
print(', '.join([city.title(),state]))
# END SOLUTION
```

Output:
```
irVinE
Irvine
IRVINE
IRvINe
arVanE
Irvine, CA
Irvine, CA
```

## Exercise: Math Module

Follow the instructions in the following cell. *Be sure to spell all variable names correctly.*

```python
# Import the math module
# BEGIN SOLUTION
import math
# END SOLUTION

# Use the function `math.pow()` to create a variable called `cube_root` that is equal to the cube root
#    of 2. I.e. 2^(1/3). Print the value of the variable `cube_root`.
# BEGIN SOLUTION
cube_root = math.pow(2,1/3)
print(cube_root)
# END SOLUTION

# Use math.pi and the function `math.sin()` to create a variable called `trig` that is equal to the value 
#    of sin(3π/2). Print the value of the variable `trig`.
# BEGIN SOLUTION
trig = math.sin(3*math.pi/2)
print(trig)
# END SOLUTION
```

Output:
```
1.2599210498948732
-1.0
```

<!-- BEGIN QUESTION -->

## Exercise: Markdown practice

In the next cell, write three sentences in the next cell on *literally any* topic that you like. The point of this is just to practice editing a markdown cell.

_Type your answer here, replacing this text._

This is a complete sentence. This is another complete sentence. An incomplete sentence.

<!-- END QUESTION -->



Tests for q1 (points: 5)
```
>>> assert city == 'irVinE'
```
```
>>> assert city_title == city.title()
```
```
>>> assert city_upper == city.upper()
```
```
>>> assert city_swapped == city.swapcase()
```
```
>>> assert city_replaced == city.replace('i', 'a')
```


Tests for q2 (points: 2)
```
>>> assert cube_root == math.pow(2, 1 / 3)
```
```
>>> assert trig == math.sin(3 * math.pi / 2)
```
