```python
# Initialize Otter
import otter
grader = otter.Notebook("Econ126_Homework_02.ipynb")
```

If the previous cell returns an error, then you need to install Otter-Grader by running the following command in a command prompt or in a Notebook cell by itself:

    pip install otter-grader

After installing, delete the cell where you ran the code to install Otter-Grader (click just to the right of the cell and type "x" or use the scissor button in the menu). Restart the Notebook using the refresh button in the menu.

# Homework 2 

**Instructions:** Complete the notebook below. Upload to completed file to the corresponding assignment on Gradescope.

## Exercise: NumPy Arrays

Follow the instructions in the following cell(s). *Be sure to spell all variable names correctly.*

```python
# Import numpy and abbreviate the namespace as np as usual
...
```

```python
# Use the `np.arange()` function to create a variable called `numbers1` that stores the integers 
#     1 through (and including) 10
...

# Print the value of `numbers1`
...
```

```python
# Use the `np.arange()` function to create a variable called `numbers2' that stores the numbers 
#     0 through (and including) 1 with a step increment of 0.01
...

# Print the value of the variable `numbers2`
...
```

```python
# Create a variable called `fifth` that is equal to the 5th value of the array `numbers2`. 
#    (Remember that the index starts counting at 0). Print the value of the variable `fifth`.

...
```

```python
# Create a variable called `last` that is equal to the last value of `numbers2`. Print the 
#    value of the variable `last`.

...
```

```python
# Create a variable called `first_twelve` that is equal to the first twelve values value of `numbers2`. Print the 
#    value of the variable `first_twelve`.

...
```

```python
# Create a variable called `last_eight` that is equal to the final eight values value of `numbers2`. Print the 
#    value of the variable `last_eight`.

...
```

```python
# Use the `np.zeros()` function to create a variable called `zeros` that is an array of 20 zeros.
...

# Change the second value of `zeros` to 1 and print 
...

...
```

## Exercise: Random Numbers

Follow the instructions in the following cell(s). *Be sure to spell all variable names correctly.*

```python
# Set the seed of NumPy's random number generator to 126
...

# Print the value of epsilon
...
```

```python
# Create a variable called `epsilon_mean` that is equal to the mean of `epsilon`. Print the 
#    value of the variable `epsilon_mean`.
...
```

```python
# Create a variable called `epsilon_std` that is equal to the standard deviation of `epsilon`. Print the 
#    value of the variable `epsilon_std`.
...
```

```python
grader.check("q_random")
```

## Exercise: The Cobb-Douglas Production Function

The Cobb-Douglas production function can be written in per worker terms as :
  \begin{align}
  y & = A k^{\alpha},
  \end{align}
where $y$ denotes output per worker, $k$ denotes capital per worker, and $A$ denotes total factor productivity or technology.

<!-- BEGIN QUESTION -->

### Part (a)

On a single axis: plot the Cobb-Douglas production for $A$ = 0.8, 1, 1.2, and 1.4 with $\alpha$ = 0.35 and $k$ ranging from 0 to 10. Each line should have a different color. Your plot must have a title and axis labels. The plot should also contain a legend that clearly indicates which line is associated with which value of $A$ and does not cover the plotted lines.

```python
# Import the pyplot module from Matplotlib as plt
...

# Select the Matlplotlib style sheet to use (Optional)
...

# Set the plotting background to white (Optional)
...
```

```python
# Set capital share (alpha)
...

# Plot production function for each of the given values for A
...

# Add x- and y-axis labels
...

# Add a title to the plot
...

# Create a legend
...

# Add a grid (Optional)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Part (b) 

**Question:** *Briefly* explain in words how increasing $A$ in Part (a) affects the shape of the production function.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Part (c)

On a single axis: plot the Cobb-Douglas production for $\alpha$ = 0.1, 0.2, 0.3, 0.4, and 0.5 with $A$ = 1 and $k$ ranging from 0 to 10. Each line should have a different color. Your plot must have a title and axis labels. The plot should also contain a legend that clearly indicates which line is associated with which value of $\alpha$ and does not cover the plotted lines.

```python
# Set TFP (A)
...

# Plot production function for each of the given values for alpha
...

# Add x- and y-axis labels
...

# Add a title to the plot
...

# Create a legend
...

# Add a grid (Optional)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Part (d)

**Question:** Explain in words how increasing $\alpha$ in Part (c) affects the shape of the production function.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Exercise: Unconstrained optimization

Consider the quadratic function:
    
\begin{align}
f(x) & = -7x^2 + 930x + 30
\end{align}
    
You will use analytic (i.e., pencil and paper) and numerical methods to find the the value of $x$ that maximizes $f(x)$. Another name for $x$ that maximizes $f(x)$ is the *argument of the maximum* value $f(x)$.

### Part (a): Analytic solution

Use standard calculus methods to solve for the value $x^*$ that maximizes $f(x)$ to **five decimal places**. In the next cell create a variable, called `xstar_exact` that is equal the value that you found for $x^*$. *For example*, if you find $x^* = 8.75359$, then enter:

    xstar_exact = 8.75359

```python
# Create a variable named `xstar_exact` equal to the value that you computed for x*. Print the value 
#    of `xstar_exact`.
...
```

```python
grader.check("q_maximization_a")
```

### Part (b):  Numerical solution

In the cells below, you will use NumPy to try to compute the argument of the maximum of $f(x)$.

```python
# Use np.arange to create a variable called `x` that is equal to the numbers 0 through 100 
# with a spacing between numbers of 0.1
...

# Create a variable called `f` that equals f(x) at each value of the array `x` just defined
...

# Use np.argmax to create a variable called `xstar_1` equal to the value of `x` that maximizes the function f(x).
...

# Print the value of `xstar_1`
...
```

```python
# Use np.arange to create a variable called `x` that is equal to the numbers 0 through 100 
# with a spacing between numbers of 0.001
...

# Create a variable called `f` that equals f(x) at each value of the array `x` just defined
...

# Use np.argmax to create a variable called `xstar_2` equal to the value of `x` that maximizes the function f(x).
...

# Print the value of `xstar_2`
...
```

```python
# Use np.arange to create a variable called `x` that is equal to the numbers 0 through *50*
# with a spacing between numbers of 0.001
...

# Create a variable called `f` that equals f(x) at each value of the array `x` just defined
...

# Use np.argmax to create a variable called `xstar_3` equal to the value of `x` that maximizes the function f(x).
...

# Print the value of `xstar_3`
...
```

<!-- BEGIN QUESTION -->

### Part (c): Evaluation

Provide answers to the follow questions in the next cell.

**Questions**

1. How did the choice of step size in the array `x` affect the accuracy of the computed results in the first two cells of Part (b)?

2. What do you think is the drawback to decreasing the stepsize in `x`?

3. In the previous cell, why did NumPy return value for `xstar` that is so different from the solution you derived in Part (a)?

*Note*: You can create enumerated lists in markdown by formating your answers like this:

    1. Answer the first question
    2. Answer the second question
    3. Answer the third question

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Exercise: Utility Maximization

Recall the two good utility maximization problem from microeconomics. Let $x$ and $y$ denote the amount of two goods that a person consumes. The person receives utility from consumption given by:
  \begin{align}
  u(x,y) & = x^{\alpha}y^{\beta} 
  \end{align}
The person has income $M$ to spend on the two goods and the price of the goods are $p_x$ and $p_y$. The consumer's budget constraint is:
  \begin{align}
  M & = p_x x + p_y y
  \end{align}
Suppose that $M = 100$, $\alpha=0.25$, $\beta=0.75$, $p_x = 1$. and $p_y = 0.5$. The consumer's problem is to maximize their utility subject to the budget constraint. While this problem can easily be solved by hand, we're going to use a computational approach. You can also solve the problem by hand to verify your solution.

<!-- BEGIN QUESTION -->

### Part (a)

Use the budget constraint to solve for $y$ in terms of $x$, $p_x$, $p_y$, and $M$. Use the result to write the consumer's utility as a function of $x$ only. Use `np.arange()` to create a variable called `x` equal to an array of values from 0 to 80 with step size equal to 0.001. Make a variable called `utility` equal to the consumer's utility. Plot the consumer's utility against $x$.

```python
# Assign values to the constants alpha, beta, M, px, py
...

# Create an array of x values between 0 and 80 with step size 0.001
...

# Create an array of utility values
...

# Plot utility against x
...

# x- and y-axis labels
...

# Title
...

# Add grid (Optional)
...
```

<!-- END QUESTION -->

### Part (b)

The NumPy function `np.max()` returns the highest value in an array and `np.argmax()` returns the value of the *index* (i.e., the location in the array) of the highest value.

```python
# Create variable `utility_max` equal to the largest value in the `utility` array.
...

# Create variable `utility_argmax` equal to index of largest value in the `utility` array.
...

# Print the value of utility_max
...


# Print the value of utility_argmax
...
```

```python
# Create variable `xstar` equal to value in `x` that maximizes utility. Note that the location in `x` will
#    be the same as that which maximizes `utility` in the previous cell
...

# Print value of `xstar`
...
```

```python
# Create variable `ystar` equal to value in `y` that maximizes utility. You can find this my solving the
#    budget constraint for y
...

# Print value of `ystar`
...
```

Tests for q_maximization_a (points: 2)
```
>>> assert np.isclose(xstar_exact, 930 / 14, atol=0.001)
```
failure_message: Incorrect solution. Check your work.


Tests for q_maximization_b (points: 2)


Tests for q_numpy (points: 5)


Tests for q_random (points: 5)
```
>>> np.random.seed(126)
>>> epsilon_answer = np.random.normal(loc=4, scale=2, size=25)
>>> assert np.allclose(epsilon, epsilon_answer)
```
failure_message: Did not create array epsilon correctly.
success_message: epsilon array created correctly.


Tests for q_utility_max_a (points: 3)
